# CS 3110/5110: Data Privacy
## Homework 8

```python
# Load the data and libraries
import pandas as pd
import numpy as np
import random
from scipy import stats
import matplotlib.pyplot as plt

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def laplace_mech_vec(vec, sensitivity, epsilon):
    return [v + np.random.laplace(loc=0, scale=sensitivity / epsilon) for v in vec]

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def gaussian_mech_vec(vec, sensitivity, epsilon, delta):
    return [v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)
            for v in vec]

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
```

```python
# preserves epsilon-differential privacy
def above_threshold(queries, df, T, epsilon):
    T_hat = T + np.random.laplace(loc=0, scale = 2/epsilon)
    
    for idx, q in enumerate(queries):
        nu_i = np.random.laplace(loc=0, scale = 4/epsilon)
        if q(df) + nu_i >= T_hat:
            return idx
    return -1 # the index of the last element
```

## Question 1 (20 points)

Implement a function `above_10000` that releases the **value** of the first query in a sequence of queries whose value is above 10000. Your function should have a **total** privacy cost equal to the privacy parameter $\epsilon$ passed in when it is called.

<!-- grade_id: cell-7d58af13aa2622cd, points: 0, grade: false, solution: true -->
```python
def make_query(status):
    def q(df):
        return len(df[df['Marital Status'] == status])
    return q

def above_10000(queries, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

queries = [make_query(status) for status in adult['Marital Status'].unique()]
print(f"above_10000 #1: {above_10000(queries, 100)}")
print(f"above_10000 #2: {above_10000(queries, 1)}")
print(f"above_10000 #3: {above_10000(queries, .01)}")
```

<!-- grade_id: cell-1cf50100348edbd6, points: 20, grade: true, solution: false -->
```python
# TEST CASE

results = [above_10000(queries, 1.0) for _ in range(100)]
print(np.mean(results))
assert np.mean(results) > 9900
assert np.mean(results) < 11000
```

## Question 2 (10 points)
In 2-3 sentences, argue informally (via the definition of the sparse vector technique, post-processing, and sequential composition), that your implementation of `above_10000` has a total privacy cost of $\epsilon$.

<!-- grade_id: cell-3d70d4d8ee9868e9, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3 (20 points)

Implement a function `bounded_all_above_10000` that releases the **value** of **$c$ queries** in a sequence of queries whose value is above 10000 (where $c$ is an analyst-provided parameter limiting the number of returned results). Your function should have a **total privacy cost** bounded by its parameter $\epsilon$.

<!-- grade_id: cell-b0f7fdabc28e37df, points: 0, grade: false, solution: true -->
```python
def bounded_all_above_10000(queries, c, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

queries = [make_query(status) for status in adult['Marital Status'].unique()]
print(f"bounded_all_above_10000 #1: {bounded_all_above_10000(queries, 3, 100)}")
print(f"bounded_all_above_10000 #2: {bounded_all_above_10000(queries, 3, 1)}")
print(f"bounded_all_above_10000 #3: {bounded_all_above_10000(queries, 3, .01)}")
```

<!-- grade_id: cell-c27f3306db2df4fa, points: 20, grade: true, solution: false -->
```python
# TEST CASE

results = [bounded_all_above_10000(queries, 2, 1.0) for _ in range(100)]
results_1 = [r[0] for r in results]
results_2 = [r[1] for r in results]

assert np.mean(results_1) > 9900
assert np.mean(results_1) < 11000
assert np.mean(results_2) > 14000
assert np.mean(results_2) < 15500
```

## Question 4 (10 points)

In 2-3 sentences, argue informally that your implementation of `bounded_all_above_10000` has privacy cost bounded by $\epsilon$.

<!-- grade_id: cell-0908f5d27b9661e6, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 5 (30 points)

Implement a function `mean_age` that computes the mean age of participants in the `adult_data` dataset. Your function should have a **total** privacy cost of $\epsilon$. It should work as follows:

1. Compute an upper clipping parameter based on the data using a differentially private algorithm
2. Clip the data using the clipping parameter
3. Use `laplace_mech` to release a differentially private mean of the clipped data

*Hint*: Use the sparse vector technique (`above_threshold`) to compute the clipping parameter. Consider using a sequence of queries that looks like `df.clip(lower=0, upper=b).sum() - df.clip(lower=0, upper=b+1).sum()`.

*Hint*: Be careful of sensitivities and set the scale of the noise accordingly!

<!-- grade_id: cell-fefcc1a8cd056ff5, points: 0, grade: false, solution: true -->
```python
bs = list(range(0, 200, 10))
df = adult['Age']
b_lower = 0

def make_query(b):
    def q(df):
        # YOUR CODE HERE
        raise NotImplementedError()
    return q

def mean_age(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()
    
for epsilon in [0.001, 0.01, 0.1, 0.5, 1, 10]:
    print(f"epsilon: {epsilon}, mean age: {mean_age(epsilon)}")
```

<!-- grade_id: cell-efca6bec52b4b3cc, points: 30, grade: true, solution: false -->
```python
# TEST CASE
results = [mean_age(1.0) for _ in range(100)]
assert np.mean(results) > 38
assert np.mean(results) < 39
```

## Question 6 (10 points)

In 3-5 sentences, describe your approach to implementing `mean_age` and argue informally that your implementation has privacy cost bounded by $\epsilon$.

<!-- grade_id: cell-c96a07678fe837cb, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 7 (20 points)

Write an algorithm to compute the (approximate) maximum of a given column of the adult dataset. Your solution should satisfy `epsilon` differential privacy.

*Hint*: you might want to use AboveThreshold; you might need to test several options for building your queries and for the threshold

<!-- grade_id: cell-de5a77cbd2059908, points: 0, grade: false, solution: true -->
```python
def max_of_col(column, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

max_of_col('Age', 1.0)
```

<!-- grade_id: cell-d325061167621b36, points: 20, grade: true, solution: false -->
```python
# TEST CASE
cols = ['Age', 'Capital Gain', 'fnlwgt']

for c in cols:
    true_val = adult[c].max()
    trials = [max_of_col(c, 10.0) for _ in range(20)]
    errors = [pct_error(true_val, t) for t in trials]
    print('Median error for column ' + c + ':', np.median(errors))
    assert np.median(errors) > 0
    assert np.median(errors) < 100
```

## Question 8 (10 points)

In 3-5 sentences, describe your approach to implementing `max_of_col` and argue that your approach has total privacy cost bounded by $\epsilon$.

<!-- grade_id: cell-13ecfdfe403adada, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE