# CS 3110/5110: Data Privacy
## Homework 5

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def laplace_mech_vec(qs, sensitivity, epsilon):
    return [laplace_mech(q, sensitivity, epsilon) for q in qs]

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def gaussian_mech_vec(qs, sensitivity, epsilon, delta):
    return [gaussian_mech(q, sensitivity, epsilon, delta) for q in qs]

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
```

## Question 1 (5 points)

Implement the `dp_occupation_histogram` function below. It should return a differentially private histogram over the `Occupation` column in the `adult` dataset. Your function should have a total privacy cost of `epsilon` and should use parallel composition.

<!-- grade_id: cell-7b8cd895f25e6d77, points: 0, grade: false, solution: true -->
```python
def dp_occupation_histogram(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

dp_occupation_histogram(1.0)
```

<!-- grade_id: cell-e017e6677a43b790, points: 5, grade: true, solution: false -->
```python
# TEST CASE
hist = adult['Occupation'].value_counts()
dp_results = [np.abs(dp_occupation_histogram(1.0)-hist) for _ in range(200)]
all_dp_results = [item for sublist in dp_results for item in sublist]
spec = [np.random.laplace(loc=0, scale=1/1.0) for _ in range(2000)]
assert stats.wasserstein_distance(all_dp_results, spec) < 2
assert stats.wasserstein_distance(all_dp_results, spec) > 0

dp_results = [np.abs(dp_occupation_histogram(0.01)-hist) for _ in range(200)]
all_dp_results = [item for sublist in dp_results for item in sublist]
spec = [np.random.laplace(loc=0, scale=1/0.01) for _ in range(2000)]
assert stats.wasserstein_distance(all_dp_results, spec) < 150
assert stats.wasserstein_distance(all_dp_results, spec) > 50
```

## Question 2 (10 points)

In 2-5 sentences each, answer the following:

- Argue that your solution for `dp_occupation_histogram` satisfies `epsilon`-differential privacy.
- How would your results change (in terms of accuracy) if you used sequential composition instead of parallel composition to implement `dp_occupation_histogram`, and why?
- If you partition the data by *both* occupation and age (i.e. a contingency table), would parallel composition still apply? Why or why not?
- If you partition the data by both occupation and age, what would happen to accuracy, and why?

<!-- grade_id: cell-dbc4948a9f3092b7, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3 (5 points)

Implement the `dp_3d_histogram` function below. It should return a differentially private 3-dimensional histogram over the `Occupation`, `Education`, and `Sex` columns (in that order) in the `adult` dataset (i.e. one count for each combination of occupation, education, and sex). As in Question 1, your function should have a total privacy cost of `epsilon` and should use parallel composition.

*Hint*: `value_counts` works on multiple columns.

<!-- grade_id: cell-bc54c2601731bf44, points: 0, grade: false, solution: true -->
```python
def dp_3d_histogram(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

dp_3d_histogram(1.0)
```

<!-- grade_id: cell-c1cd029d467d2b7a, points: 5, grade: true, solution: false -->
```python
# TEST CASE
hist = adult[['Occupation', 'Education', 'Sex']].value_counts()
dp_results = [(dp_3d_histogram(1.0)-hist).dropna() for _ in range(50)]
all_dp_results = [item for sublist in dp_results for item in sublist]
spec = [np.random.laplace(loc=0, scale=1/1.0) for _ in range(2000)]
assert stats.wasserstein_distance(all_dp_results, spec) < 0.3
assert stats.wasserstein_distance(all_dp_results, spec) > 0

dp_results = [(dp_3d_histogram(0.1)-hist).dropna() for _ in range(50)]
all_dp_results = [item for sublist in dp_results for item in sublist]
spec = [np.random.laplace(loc=0, scale=1/0.1) for _ in range(2000)]
assert stats.wasserstein_distance(all_dp_results, spec) < 3
assert stats.wasserstein_distance(all_dp_results, spec) > 0
```

## Question 4 (10 points)

In 2-5 sentences each, answer the following:

- Argue that your solution for `dp_3d_histogram` satisfies `epsilon`-differential privacy.
- Which histogram do you expect to have higher relative error (percent error): the 1-dimensional one from Question 1, or the 3-dimensional one from Question 3? Why?

<!-- grade_id: cell-4879140d72a33f07, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 5 (5 points)

Implement the `dp_three_histograms` function below. It should return a tuple of three differentially private histograms:

1. A 1-dimensional histogram over the `Occupation` column
2. A 2-dimensional histogram over the `Occupation` and `Education` columns
3. A 3-dimensional histogram over the `Occupation`, `Education`, and `Sex` columns

Your function should have a total privacy cost of `epsilon`.

Hint: take care with composition!

<!-- grade_id: cell-80b3ad750d09e310, points: 0, grade: false, solution: true -->
```python
def dp_three_histograms(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

dp_three_histograms(1.0)
```

<!-- grade_id: cell-cdf8dca2564c2251, points: 5, grade: true, solution: false -->
```python
# TEST CASE
hists = [adult['Occupation'].value_counts(),
         adult[['Occupation', 'Education']].value_counts(),
         adult[['Occupation', 'Education', 'Sex']].value_counts()]
dp_results = [dp_three_histograms(1.0) for _ in range(100)]
assert all(len(r) == 3 for r in dp_results)

for i, hist in enumerate(hists):
    all_dp_results = [item for r in dp_results for item in (r[i]-hist).dropna()]
    spec = [np.random.laplace(loc=0, scale=3/1.0) for _ in range(2000)]
    assert stats.wasserstein_distance(all_dp_results, spec) < 1
    assert stats.wasserstein_distance(all_dp_results, spec) > 0
```

## Question 6 (5 points)

Using the 3-point structure described in lecture, argue that your solution for `dp_three_histograms` satisfies `epsilon`-differential privacy. Be precise about composition.

<!-- grade_id: cell-8d22b716c8af7acf, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE