# CS 3110/5110: Data Privacy
## Homework 6

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
retail = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/retail.csv').dropna()
```

## Question 1 (10 points)

Implement the function `adult_hists`, which should return three histograms based on the `adult` dataset:

1. A 1-d histogram of the `Occupation` column
2. A 1-d histogram of the `Education` column
3. A 2-d histogram of the `Race` and `Sex` columns

Your function should have a **total** privacy cost of `epsilon` and a **person-level** unit of privacy.

<!-- grade_id: cell-63f79c9bbc970326, points: 0, grade: false, solution: true -->
```python
def adult_hists(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-09e957f0bec6bb1c, points: 10, grade: true, solution: false -->
```python
# TEST CASE
occupation, education, multi = adult_hists(1.0)
assert occupation.sum() > 30000
assert education.sum() > 32000
assert multi.sum() > 32000
```

## Question 2 (10 points)

Argue that your implementation of `adult_hists` has a total privacy cost of `epsilon` and a person-level unit of privacy.

<!-- grade_id: cell-5014a123a618927a, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3 (10 points)

Implement the function `retail_hist`, which should return one histogram based on the `retail` dataset:

1. A 1-d histogram of the `Country` column

Your function should have a **total** privacy cost of `epsilon` and an **event-level** unit of privacy.

<!-- grade_id: cell-c3574b3bf92edbb9, points: 0, grade: false, solution: true -->
```python
def retail_hist(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-42a43ac58be658e0, points: 10, grade: true, solution: false -->
```python
# TEST CASE
country_hist = retail_hist(1.0)
assert country_hist.sum() > 7400
```

## Question 3 (10 points)

Argue that your implementation of `retail_hist` has a total privacy cost of `epsilon` and an event-level unit of privacy.

<!-- grade_id: cell-5a28e9b46d3e9bb2, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 4 (10 points)

Implement the function `retail_hist_bounded`, which returns a 1-d histogram of the `Country` column, with a total privacy cost of `epsilon`, and with a *user-level* unit of privacy. Use contribution bounding to enforce the unit of privacy.

<!-- grade_id: cell-76afc774a939561f, points: 0, grade: false, solution: true -->
```python
def retail_hist_bounded(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-bc3ec8143e03943e, points: 10, grade: true, solution: false -->
```python
# TEST CASE
country_hist = retail_hist_bounded(1.0)
assert country_hist.sum() > 6000
```

## Question 5 (10 points)

Argue that your implementation of `retail_hist_bounded` has a total privacy cost of `epsilon` and a person-level unit of privacy.

<!-- grade_id: cell-5a28e9b46d3a9bb2, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 6 (20 points)

Implement `retail_avg_quantity`, a function that returns the average quantity of orders in the `retail` dataset, with a total privacy cost of `epsilon` and a person-level unit of privacy.

<!-- grade_id: cell-4a82c701b4af2636, points: 0, grade: false, solution: true -->
```python
def retail_avg_quantity(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-de21931b1bfb3642, points: 20, grade: true, solution: false -->
```python
# TEST CASE
avg = retail_avg_quantity(1)
assert avg > 7
assert avg < 13
```

## Question 7 (10 points)

Argue that your implementation of `retail_avg_quantity` has a total privacy cost of `epsilon` and a person-level unit of privacy.

<!-- grade_id: cell-5a28e9b46d3f9bb2, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 8 (10 points)

In 1-3 sentences each, answer the following:

- What impact did the choice of contribution bound have on the average?
- How did you arrive at your choice of contribution bound?
- What is the better choice of privacy unit for the `retail` dataset: event-level or person-level?
- What is an example of a potential negative outcome associated with choosing the event-level unit of privacy for this dataset?

<!-- grade_id: cell-f4153a007ae8b69d, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE