# CS 3110/5110: Data Privacy
## In-Class Exercise, week of 10/13/2025

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
retail = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/retail.csv')
```

```python
def gaussian_mech_RDP(val, sensitivity, alpha, epsilon):
    sigma = np.sqrt((sensitivity**2 * alpha) / (2 * epsilon))
    return val + np.random.normal(loc=0, scale=sigma)
```

## Question 1

Consider the retail dataset and the following counting query on it. What is the query's sensitivity?

```python
def retail_query():
    return len(retail[retail['CustomerID'] == 12748])

retail_query()
```

<!-- grade_id: cell-619da35fc3dabdc8, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2

What are we actually protecting in the retail dataset, when the unit of privacy is set to:

1. Event-level / row-level
2. Person-level / user-level

<!-- grade_id: cell-96f996a56ca77090, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3

What is the *effective* person-level $\epsilon$ value for the following query, which assumes a row-level unit of privacy?

```python
def dp_retail_query():
    result = retail_query()
    return laplace_mech(result, sensitivity=1, epsilon=1)

dp_retail_query()
```

<!-- grade_id: cell-76d3e9d0559d8834, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 4

Use *contribution bounding* to adjust the unit of privacy to be person-level privacy, and implement a differentially private function to return the orders originating from the United Kingdom.

<!-- grade_id: cell-3e24aaa7cc0a4d79, points: 0, grade: false, solution: true -->
```python
def bound_contributions(bound):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def dp_retail_query(bound, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()
    
dp_retail_query(100, 1.0)
```

<!-- grade_id: cell-6e12c0f1e20ad0f1, points: 0, grade: true, solution: false -->
```python
# TEST CASE

results = [dp_retail_query(100, 1.0) for _ in range(100)]
errors = [pct_error(len(retail[retail['Country'] == 'United Kingdom']), r) for r in results]
print('mean error:', np.mean(errors))

assert np.mean(errors) > 20
assert np.mean(errors) < 50
```

## Question 5

Write a function to calculate the L1 **local sensitivity** of the following mean query.

```python
def mean_query():
    b = 125
    return adult['Age'][adult['Education'] == 'Bachelors'].clip(lower=0, upper=b).mean()

mean_query()
```

<!-- grade_id: cell-f39828c2d85af776, points: 0, grade: false, solution: true -->
```python
def mean_query_local_sensitivity():
    # YOUR CODE HERE
    raise NotImplementedError()

mean_query_local_sensitivity()
```

<!-- grade_id: cell-f1bb755347650609, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert (mean_query_local_sensitivity() - 0.02333831217326363) < 0.0001
```

## Question 6

Why can't we use the local sensitivity from question 1 directly to achieve differential privacy for `mean_query`?

<!-- grade_id: cell-55d134a1eb26aec9, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 7

Calculate the empirical *cumulative distribution function* of ages in the adult dataset. Use the RDP Gaussian mechanism.

<!-- grade_id: cell-ce242a7d80d36e05, points: 5, grade: true, solution: true -->
```python
def calculate_cdf():
    alpha = 5
    epsilon = 0.0001
    
    # YOUR CODE HERE
    raise NotImplementedError()

cdf_vals = calculate_cdf()
plt.plot(range(20, 80), cdf_vals);
```

## Question 8

What is the *total privacy cost* in RDP of your solution above?

<!-- grade_id: cell-234f1767c58fcaff, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 9

What is the *total privacy cost* in $(\epsilon, \delta)$-differential privacy of your solution above, for $\delta = 10^{-5}$?

<!-- grade_id: cell-f492eea07aa320fc, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 10

Repeat the earlier question, but for zero-concentrated differential privacy.

<!-- grade_id: cell-1854b003d5ab2ebf, points: 5, grade: true, solution: true -->
```python
def gaussian_mech_zCDP(val, sensitivity, rho):
    # YOUR CODE HERE
    raise NotImplementedError()

def calculate_cdf_zcdp():
    rho = 0.00001
    
    # YOUR CODE HERE
    raise NotImplementedError()

cdf_vals = calculate_cdf_zcdp()
plt.plot(range(20, 80), cdf_vals);
```

```python

```