# CS 3110/5110: Data Privacy
## In-Class Exercise, week of 10/16/2023

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
```

```python
def gaussian_mech_RDP(val, sensitivity, alpha, epsilon):
    sigma = np.sqrt((sensitivity**2 * alpha) / (2 * epsilon))
    return val + np.random.normal(loc=0, scale=sigma)
```

## Question 1

Write a function to calculate the L1 **local sensitivity** of the following mean query.

```python
def mean_query():
    b = 125
    return adult['Age'][adult['Education'] == 'Bachelors'].clip(lower=0, upper=b).mean()

mean_query()
```

<!-- grade_id: cell-f39828c2d85af776, points: 0, grade: false, solution: true -->
```python
def mean_query_local_sensitivity():
    # YOUR CODE HERE
    raise NotImplementedError()

mean_query_local_sensitivity()
```

<!-- grade_id: cell-f1bb755347650609, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert (mean_query_local_sensitivity() - 0.02333831217326363) < 0.0001
```

## Question 2

Why can't we use the local sensitivity from question 1 directly to achieve differential privacy for `mean_query`?

<!-- grade_id: cell-55d134a1eb26aec9, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3

Implement an upper bound on $A(f, x, k)$, the local sensitivity of a mean query at distance $k$.

<!-- grade_id: cell-bfe45a4f5261718f, points: 0, grade: false, solution: true -->
```python
def ls_mean_at_distance(df, b, k):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e19c35e52e3a15c0, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert ls_mean_at_distance(adult['Age'], 125, 500) == 0.0038988178784192633
assert ls_mean_at_distance(adult['Age'], 125, 5000) == 0.004535394216465295
```

## Question 4

Implement an upper bound on $D(f, x, s)$, the distance from $x$ to a database where the local sensitivity of the mean query is greater than or equal to $s$.

<!-- grade_id: cell-2750bb4f33a0b420, points: 0, grade: false, solution: true -->
```python
def dist_to_high_ls_mean(df, b, s):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-77a206124ddb3baf, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert dist_to_high_ls_mean(adult['Age'], 125, 0.1) == 31311 
assert dist_to_high_ls_mean(adult['Age'], 125, 0.01) == 20061
assert dist_to_high_ls_mean(adult['Age'], 125, 0.001) == 0
```

## Question 5

Use the Propose-Test-Release (PTR) framework to calculate the mean age in the adult dataset with differential privacy.

<!-- grade_id: cell-62e95c3348999e9a, points: 0, grade: false, solution: true -->
```python
def mean_ptr(epsilon, delta):
    # YOUR CODE HERE
    raise NotImplementedError()

mean_ptr(1.0, 10e-5)
```

<!-- grade_id: cell-efb8d6da908b8ae9, points: 5, grade: true, solution: false -->
```python
# TEST CASE

true_answer = adult['Age'].mean()
trials = [mean_ptr(1.0, 10e-5) for _ in range(20)]
errors = [pct_error(true_answer, x) for x in trials]
assert np.mean(errors) < 1.0
```

## Question 6

Use the sample-and-aggregate framework to calculate the averge of ages in the adult dataset with differential privacy.

<!-- grade_id: cell-8afedb58ded17a79, points: 5, grade: true, solution: true -->
```python
def saa_avg_age(k, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

saa_avg_age(600, 1)    
```

## Question 7

What is the impact of the parameter `k`?

<!-- grade_id: cell-6f57645dec8e7707, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 8

Calculate the empirical *cumulative distribution function* of ages in the adult dataset. Use the RDP Gaussian mechanism.

<!-- grade_id: cell-ce242a7d80d36e05, points: 5, grade: true, solution: true -->
```python
def calculate_cdf():
    alpha = 5
    epsilon = 0.0001
    
    # YOUR CODE HERE
    raise NotImplementedError()

cdf_vals = calculate_cdf()
plt.plot(range(20, 80), cdf_vals);
```

## Question 9

What is the *total privacy cost* in RDP of your solution above?

<!-- grade_id: cell-234f1767c58fcaff, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 10

What is the *total privacy cost* in $(\epsilon, \delta)$-differential privacy of your solution above, for $\delta = 10^{-5}$?

<!-- grade_id: cell-f492eea07aa320fc, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 11

Repeat question 3, but for zero-concentrated differential privacy.

<!-- grade_id: cell-1854b003d5ab2ebf, points: 5, grade: true, solution: true -->
```python
def gaussian_mech_zCDP(val, sensitivity, rho):
    # YOUR CODE HERE
    raise NotImplementedError()

def calculate_cdf_zcdp():
    rho = 0.00001
    
    # YOUR CODE HERE
    raise NotImplementedError()

cdf_vals = calculate_cdf_zcdp()
plt.plot(range(20, 80), cdf_vals);
```